## Problem 4 - Nested `for` loops (*optional, does not affect grading*)

**Warning:** This is a difficult problem. Try to do this task only if you have confidence in programming, and you are up for a little challenge. :) 

In addition to having single `for` loops that iterate across some variable range, it is possible to *nest* `for` loops within one another.

Consider the example below:

```python
>>> for char in 'dog':
...     for char2 in 'cat':
...         print (char, char2)
    d c
    d a
    d t
    o c
    o a
    o t
    g c
    g a
    g t
```

Here, you can see that in the first pass through the first `for` loop, the value of `char` is `d`.
Entering the inner (or nested) loop, `char2` is set to `c`.
After this, the output is written to the screen and since there are more letters to loop over in the inner `for` loop, the value of `char2` will be updated upon the next pass.
The second time through the inner loop, `char2` is set to `a` while `char` remains `d`.
Like this, the inner loop will run through all of the letters in `cat` for each time that `char` is updated in the outer loop.
It doesn't take too much imagination to figure out this is a very useful concept.

For this problem you should do following:

1. Create a variable `star` with text `"*"` and an empty string variable `text`. Recall, you can created empty string variables by assigning `""` as their value.
2. Use nested `for` loops and the variables above to produce the text formation below when `print(text)` is run at the end of your script.

    ```
    *******
    *******
    *******
    ```
3. Create a variable `line` with text `"-"` and an empty string variable `flag`.
4. Use nested `for` loops and the variables above to produce the text formation below when `print(flag)` is run at the end of your script. **Note**: You will need to use conditional statements to produce the desired output.

    ```
    *******------------
    *******------------
    *******------------
    -------------------
    -------------------
    ```


### Part 1

- Your code for producing the asterisk square:

<!-- grade_id: asterix-square, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
star = '*'
text = ""

for i in range(3):
    for j in range(7):
        text+= star
    text += '\n'

### END SOLUTION
text
```

<!-- grade_id: asterix-square-test, points: 0, grade: true, solution: false -->
```python
print(text)
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
correct_shape = '*******\n*******\n*******\n'
assert_equal(text, correct_shape)
### END HIDDEN TESTS
```

### Part 2

- Your code for producing the flag:

<!-- grade_id: flag, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
star = '*'
flag = ""
line = '-'

for i in range(5):
    # For the first three lines add asterix and lines
    if i < 3:
        for j in range(7):
            flag+= star
            
        for k in range(12):
            flag+=line
    # For the last two lines add only lines
    else:
        for k in range(19):
            flag+=line
    flag += '\n'
### END SOLUTION
```

<!-- grade_id: flag_test, points: 0, grade: true, solution: false -->
```python
print(flag)
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
correct_shape = '*******------------\n*******------------\n*******------------\n-------------------\n-------------------\n'
assert_equal(flag, correct_shape)
### END HIDDEN TESTS
```

### All done!

Congrats, if you are here, you have finished all assignments for week 3, even the extra one! Bad-ass. &#128526;