## Problem 2 - Classifying temperatures (**8 points**)

### Overview

This problem is meant to introduce you to a very commonly used and useful concept of data classification.
In this problem your aim is to classify daily temperatures stored in `temperatures` list into four different classes:

  1. **Cold** ==> temperatures below -2 degrees (Celsius)
  2. **Slippery** ==> temperatures warmer or equal to -2 degrees and up to +2 degrees (Celsius)
  3. **Comfortable** ==> temperatures warmer or equal to +2 degrees and up to +15 degrees (Celsius)
  4. **Warm** ==> temperatures warmer or equal to +15 degrees (Celsius)

To solve this problem, you should modify and fill in the missing parts in the following cells.
In total, there are three tasks that you should solve according the directions. In addition, there is
one additional task (Task 4) that is **optional** for the students that are quicker (*does not affect grading*).

**Notice**: follow closely the instructions, e.g. using **exactly** the same variable names as mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting as in the instructions. 

**Your score on this problem will be based on following criteria:**

 - Using for loop to iterate over the temperature values
 - Using conditional statements to find out if a value is within certain value range
 - Printing information for the user
 - Including comments that explain what most lines in the code do
 - Uploading your Notebook to your GitHub repository for this week's Exercise.

### Data

```python
# A list of night-time, day-time and evening temperatures for April 2013 measured in Helsinki Malmi Airport. 
# The list contains 90 values since there are 3 values for each day (April has 30 days).
# The first value of a single day represents night, the second one day-time, and the third one evening temperature.

temperatures = [-5.4, 1.0, -1.3, -4.8, 3.9, 0.1, -4.4, 4.0, -2.2, -3.9, 4.4,
                -2.5, -4.6, 5.1, 2.1, -2.4, 1.9, -3.3, -4.8, 1.0, -0.8, -2.8,
                -0.1, -4.7, -5.6, 2.6, -2.7, -4.6, 3.4, -0.4, -0.9, 3.1, 2.4,
                1.6, 4.2, 3.5, 2.6, 3.1, 2.2, 1.8, 3.3, 1.6, 1.5, 4.7, 4.0,
                3.6, 4.9, 4.8, 5.3, 5.6, 4.1, 3.7, 7.6, 6.9, 5.1, 6.4, 3.8,
                4.0, 8.6, 4.1, 1.4, 8.9, 3.0, 1.6, 8.5, 4.7, 6.6, 8.1, 4.5,
                4.8, 11.3, 4.7, 5.2, 11.5, 6.2, 2.9, 4.3, 2.8, 2.8, 6.3, 2.6,
                -0.0, 7.3, 3.4, 4.7, 9.3, 6.4, 5.4, 7.6, 5.2]
```

1: Create four empty lists for different temperature classes, i.e.:

  - `cold`
  - `slippery`
  - `comfortable`
  - `warm`
  
Be careful to use **exactly** these names for your lists.

<!-- grade_id: temp-lists, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
cold = []
slippery = []
comfortable = []
warm = []
### END SOLUTION
```

<!-- grade_id: temp-lists-tests, points: 0.5, grade: true, solution: false -->
```python
# Test print for all lists (they should be empty at this point)
print(cold, slippery, comfortable, warm)

### BEGIN HIDDEN TESTS
from nose.tools import ok_

# Check that variables exists
ok_('cold' in locals())
ok_('slippery' in locals())
ok_('comfortable' in locals())
ok_('warm' in locals())

### END HIDDEN TESTS
```

2: Iterate over temperatures and add temperatures to different temperature classes as defined below (*5.5 points*):
  - **Cold** ==> temperatures below -2 degrees (Celsius)
  - **Slippery** ==> temperatures warmer or equal to -2 degrees and up to +2 degrees (Celsius)
  - **Comfortable** ==> temperatures warmer of equal to +2 degrees and up to +15 degrees (Celsius)
  - **Warm** ==> temperatures above or equal to +15 degrees (Celsius)

<!-- grade_id: temp-classes, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# Iterate over values using for loop
for temp in temperatures:
    # Create conditions for different temperature ranges
    if temp < -2:
        # Add to list
        cold.append(temp)
    elif temp >= -2 and temp < 2:
        slippery.append(temp)
    elif temp >= 2 and temp < 15:
        comfortable.append(temp)
    elif temp >= 15:
        warm.append(temp)
### END SOLUTION
```

<!-- grade_id: temp-classes-tests, points: 5.5, grade: true, solution: false -->
```python
# Test prints for all lists (now they should contain values)
print(cold, slippery, comfortable, warm)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
correct_cold = [-5.4, -4.8, -4.4, -2.2, -3.9, -2.5, -4.6, -2.4, -3.3, -4.8, -2.8, -4.7, -5.6, -2.7, -4.6]
correct_slippery = [1.0, -1.3, 0.1, 1.9, 1.0, -0.8, -0.1, -0.4, -0.9, 1.6, 1.8, 1.6, 1.5, 1.4, 1.6, -0.0]
correct_comfortable = [3.9, 4.0, 4.4, 5.1, 2.1, 2.6, 3.4, 3.1, 2.4, 4.2, 3.5, 2.6, 3.1, 2.2, 3.3, 4.7, 4.0, 3.6, 4.9, 4.8, 5.3, 5.6, 4.1, 3.7, 7.6, 6.9, 5.1, 6.4, 3.8, 4.0, 8.6, 4.1, 8.9, 3.0, 8.5, 4.7, 6.6, 8.1, 4.5, 4.8, 11.3, 4.7, 5.2, 11.5, 6.2, 2.9, 4.3, 2.8, 2.8, 6.3, 2.6, 7.3, 3.4, 4.7, 9.3, 6.4, 5.4, 7.6, 5.2]
correct_warm = []

# Test that values in temperature classes are correct
assert_equal(cold, correct_cold)
assert_equal(slippery, correct_slippery)
assert_equal(comfortable, correct_comfortable)
assert_equal(warm, correct_warm)

### END HIDDEN TESTS

```

3: Questions - print the answers (*1 point*)

- How many times was it slippery during the study period?

<!-- grade_id: slippery-times, points: 0, grade: false, solution: true -->
```python
# Edit the variable to find correct answer
slippery_times = 'XXX'

### BEGIN SOLUTION
slippery_times = len(slippery)
### END SOLUTION
```

<!-- grade_id: slippery-times-test, points: 0.33, grade: true, solution: false -->
```python
# Print the answer
print("In April 2013 it was slippery", slippery_times, "times.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
# It was slippery 16 times
assert_equal(slippery_times, 16)

### END HIDDEN TESTS
```


- How many times was it warm during the study period?

<!-- grade_id: warm-times, points: 0, grade: false, solution: true -->
```python
# Edit the variable to find correct answer
warm_times = 'XXX'

### BEGIN SOLUTION
warm_times = len(warm)
### END SOLUTION
```

<!-- grade_id: warm-times-test, points: 0.33, grade: true, solution: false -->
```python
# Print the answer
print("In April 2013 it was warm", warm_times, "times.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
# It was slippery 16 times
assert_equal(warm_times, 0)

### END HIDDEN TESTS
```

- How many times was it cold during the study period?

<!-- grade_id: cold-times, points: 0, grade: false, solution: true -->
```python
# Edit the variable to find correct answer
cold_times = 'XXX'

### BEGIN SOLUTION
cold_times = len(cold)
### END SOLUTION
```

<!-- grade_id: cold-times-test, points: 0.34, grade: true, solution: false -->
```python
# Print the answer
print("In April 2013 it was cold", cold_times, "times.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
# It was slippery 16 times
assert_equal(cold_times, 15)

### END HIDDEN TESTS
```

### Comment section:

Here, we ask a few questions to make sure you have understand the concepts etc. Answer shortly with a couple of sentences to following questions.

You can also write any feedback or questions concerning this problem (we will also use this space for giving feedback to you, if necessary).

#### Questions

1: Is the concept of conditional statements clear to you? If not, what is difficult to understand?

2: Did you include comments in your code blocks?
  
Write your answers below (double click to activate the cell below):

<!-- grade_id: p2-style-assessment, points: 1, grade: true, solution: true -->


#### Done!

That's it! Now you are ready to continue with [Problem 3](Exercise-3-problem-3.ipynb).