# Exercise 3

The exercise for this week is meant to help you to understand `for` loops and conditional statements in Python.
Below you have a series of problems in which you are asked to edit the notebook files and write the code necessary to produce the desired results.

After making your changes to the Notebook, you will need to upload it to GitHub.


- **Exercise 3 is due by the start of lecture on 26.9**.

- Don't forget to check out the [hints for this week's exercise](https://geo-python.github.io/2018/lessons/L3/exercise-3-hints.html) if you're having trouble.

- Scores on this exercise are out of **20 points**.

## Before starting

Test that your Python version is up-to-date. Our automatic grading system requires to have IPython version which is higher than 3. Test your version by running following cell (shift + Enter):

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it. Ask help, if you don't know how."
print("Your system is okay, you can continue! :) ")
```

## Problem 1 - Batch processing data files with a `for` loop (**4 Points**)

This problem is meant to simulate a common problem dealing with data files: batch processing.

Batch processing involves using scripts to process many data files, and one common task is generating a list of filenames that will be processed, or saved to disk.

For this problem you will need to modify the code blocks below and add codes to produce the desired outcomes. 

**Notice**: follow closely the instructions, e.g. using **exactly** the same variable names as mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting as in the instructions. 


1: Create a new variable called `basename` that contains text `"Station"`.

<!-- grade_id: station-name, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
basename = "Station"
### END SOLUTION
```

<!-- grade_id: station-name-correct, points: 0.25, grade: true, solution: false -->
```python
# Test print of the variable
print(basename)

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal
# Validate variable name
ok_('basename' in locals())

# Validate value
assert_equal(basename, 'Station')
### END HIDDEN TESTS
```

2: Create a new variable `filenames` that is an empty list.

<!-- grade_id: empty-filenames-list, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
filenames = []
### END SOLUTION
```

<!-- grade_id: empty-filenames-list-correct, points: 0.25, grade: true, solution: false -->
```python
# Test print of the variable
print(filenames)

### BEGIN HIDDEN TESTS

from nose.tools import ok_
ok_('filenames' in locals())
ok_(isinstance(filenames, list))

### END HIDDEN TESTS
```

3: Iterate over the number range 0-20 and
  - Create a variable `station` that contains the 1) text from `basename` variable, 2) the number, and 3) the file extension `.txt`  
  - Add the content of `station` to `filenames` list which should have following content in the end:

      ```
      ['Station_0.txt', 'Station_1.txt', 'Station_2.txt', 'Station_3.txt',
       'Station_4.txt', 'Station_5.txt', 'Station_6.txt', 'Station_7.txt',
       'Station_8.txt', 'Station_9.txt', 'Station_10.txt', 'Station_11.txt',
       'Station_12.txt', 'Station_13.txt', 'Station_14.txt', 'Station_15.txt',
       'Station_16.txt', 'Station_17.txt', 'Station_18.txt', 'Station_19.txt',
       'Station_20.txt']
      ```

<!-- grade_id: populate-filenames-list, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# Iterate from 0 up to 20
for i in range(21):
    # Create the name
    station = basename + "_" + str(i) + ".txt"
    # Append to lis
    filenames.append(station)

# Print the contents
print(filenames)
    
### END SOLUTION
```

<!-- grade_id: populate-filenames-list-1, points: 2.5, grade: true, solution: false -->
```python
# Test print of the results
print(filenames)

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal

# Check that 'station' variable exists
ok_('station' in locals())

# Check that the value of the last station is correct
assert_equal(station.lower().strip(), "station_20.txt")

# Check that there exists 21 values in the list
ok_(len(filenames) == 21)

correct_list = ['station_0.txt', 'station_1.txt', 'station_2.txt', 'station_3.txt', 'station_4.txt', 'station_5.txt', 'station_6.txt', 'station_7.txt', 'station_8.txt', 'station_9.txt', 'station_10.txt', 'station_11.txt', 'station_12.txt', 'station_13.txt', 'station_14.txt', 'station_15.txt', 'station_16.txt', 'station_17.txt', 'station_18.txt', 'station_19.txt', 'station_20.txt']

# Check that the contents match (use lowercase to allow some variation in upper/lower case)
for value in filenames:
    ok_(value.strip().lower() in correct_list)

### END HIDDEN TESTS
```

**Your score on this problem will be based on following criteria:**

 - Creating and using variables to produce the desired text format
 - Successfully using for loop to iterate over desired set of numbers
 - Successfully producing the desired filename
 - Including comments that explain what most lines in the code do
 - Uploading your notebook to your GitHub repository for this week's Exercise


### Comment section:

Here, we ask a few questions to make sure you have understand the concepts etc. Answer shortly with a 

You can also write any feedback or questions concerning this problem (we will also use this space for giving feedback to you, if necessary).

1: Is the concept of a loop clear to you? If not, what is difficult to understand?
     

2: Did you include comments in your code blocks?
  
Write your answers below:

<!-- grade_id: p1-style-assessment, points: 1, grade: true, solution: true -->


## Problem 2 - Classifying temperatures (**8 points**)

### Overview

This problem is meant to introduce you to a very commonly used and useful concept of data classification.
In this problem your aim is to classify daily temperatures stored in `temperatures` list into four different classes:

  1. **Cold** ==> temperatures below -2 degrees (Celsius)
  2. **Slippery** ==> temperatures warmer or equal to -2 degrees and up to +2 degrees (Celsius)
  3. **Comfortable** ==> temperatures warmer or equal to +2 degrees and up to +15 degrees (Celsius)
  4. **Warm** ==> temperatures warmer or equal to +15 degrees (Celsius)

To solve this problem, you should modify and fill in the missing parts in the following cells.
In total, there are three tasks that you should solve according the directions. In addition, there is
one additional task (Task 4) that is **optional** for the students that are quicker (*does not affect grading*).

**Your score on this problem will be based on following criteria:**

 - Using for loop to iterate over the temperature values
 - Using conditional statements to find out if a value is within certain value range
 - Printing information for the user
 - Including comments that explain what most lines in the code do
 - Uploading your Notebook to your GitHub repository for this week's Exercise.



```python
# A list of night-time, day-time and evening temperatures for April 2013 measured in Helsinki Malmi Airport. 
# The list contains 90 values since there are 3 values for each day (April has 30 days).
# The first value of a single day represents night, the second one day-time, and the third one evening temperature.

temperatures = [-5.4, 1.0, -1.3, -4.8, 3.9, 0.1, -4.4, 4.0, -2.2, -3.9, 4.4,
                -2.5, -4.6, 5.1, 2.1, -2.4, 1.9, -3.3, -4.8, 1.0, -0.8, -2.8,
                -0.1, -4.7, -5.6, 2.6, -2.7, -4.6, 3.4, -0.4, -0.9, 3.1, 2.4,
                1.6, 4.2, 3.5, 2.6, 3.1, 2.2, 1.8, 3.3, 1.6, 1.5, 4.7, 4.0,
                3.6, 4.9, 4.8, 5.3, 5.6, 4.1, 3.7, 7.6, 6.9, 5.1, 6.4, 3.8,
                4.0, 8.6, 4.1, 1.4, 8.9, 3.0, 1.6, 8.5, 4.7, 6.6, 8.1, 4.5,
                4.8, 11.3, 4.7, 5.2, 11.5, 6.2, 2.9, 4.3, 2.8, 2.8, 6.3, 2.6,
                -0.0, 7.3, 3.4, 4.7, 9.3, 6.4, 5.4, 7.6, 5.2]
```

1: Create four empty lists for different temperature classes, i.e.:

  - `cold`
  - `slippery`
  - `comfortable`
  - `warm`
  
Be careful to use **exactly** these names for your lists.

<!-- grade_id: temp-lists, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
cold = []
slippery = []
comfortable = []
warm = []
### END SOLUTION
```

<!-- grade_id: temp-lists-tests, points: 0.5, grade: true, solution: false -->
```python
# Test print for all lists (they should be empty at this point)
print(cold, slippery, comfortable, warm)

### BEGIN HIDDEN TESTS
from nose.tools import ok_

# Check that variables exists
ok_('cold' in locals())
ok_('slippery' in locals())
ok_('comfortable' in locals())
ok_('warm' in locals())

### END HIDDEN TESTS
```

2: Iterate over temperatures and add temperatures to different temperature classes as defined below (*5.5 points*):
  - **Cold** ==> temperatures below -2 degrees (Celsius)
  - **Slippery** ==> temperatures warmer or equal to -2 degrees and up to +2 degrees (Celsius)
  - **Comfortable** ==> temperatures warmer of equal to +2 degrees and up to +15 degrees (Celsius)
  - **Warm** ==> temperatures above or equal to +15 degrees (Celsius)

<!-- grade_id: temp-classes, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# Iterate over values using for loop
for temp in temperatures:
    # Create conditions for different temperature ranges
    if temp < -2:
        # Add to list
        cold.append(temp)
    elif temp >= -2 and temp < 2:
        slippery.append(temp)
    elif temp >= 2 and temp < 15:
        comfortable.append(temp)
    elif temp >= 15:
        warm.append(temp)
### END SOLUTION
```

<!-- grade_id: temp-classes-tests, points: 5.5, grade: true, solution: false -->
```python
# Test prints for all lists (now they should contain values)
print(cold, slippery, comfortable, warm)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
correct_cold = [-5.4, -4.8, -4.4, -2.2, -3.9, -2.5, -4.6, -2.4, -3.3, -4.8, -2.8, -4.7, -5.6, -2.7, -4.6]
correct_slippery = [1.0, -1.3, 0.1, 1.9, 1.0, -0.8, -0.1, -0.4, -0.9, 1.6, 1.8, 1.6, 1.5, 1.4, 1.6, -0.0]
correct_comfortable = [3.9, 4.0, 4.4, 5.1, 2.1, 2.6, 3.4, 3.1, 2.4, 4.2, 3.5, 2.6, 3.1, 2.2, 3.3, 4.7, 4.0, 3.6, 4.9, 4.8, 5.3, 5.6, 4.1, 3.7, 7.6, 6.9, 5.1, 6.4, 3.8, 4.0, 8.6, 4.1, 8.9, 3.0, 8.5, 4.7, 6.6, 8.1, 4.5, 4.8, 11.3, 4.7, 5.2, 11.5, 6.2, 2.9, 4.3, 2.8, 2.8, 6.3, 2.6, 7.3, 3.4, 4.7, 9.3, 6.4, 5.4, 7.6, 5.2]
correct_warm = []

# Test that values in temperature classes are correct
assert_equal(cold, correct_cold)
assert_equal(slippery, correct_slippery)
assert_equal(comfortable, correct_comfortable)
assert_equal(warm, correct_warm)

### END HIDDEN TESTS

```

3: Questions - print the answers (*1 point*)

- How many times was it slippery during the study period?

<!-- grade_id: slippery-times, points: 0, grade: false, solution: true -->
```python
# Edit the variable to find correct answer
slippery_times = 'XXX'

### BEGIN SOLUTION
slippery_times = len(slippery)
### END SOLUTION
```

<!-- grade_id: slippery-times-test, points: 0.33, grade: true, solution: false -->
```python
# Print the answer
print("In April 2013 it was slippery", slippery_times, "times.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
# It was slippery 16 times
assert_equal(slippery_times, 16)

### END HIDDEN TESTS
```


- How many times was it warm during the study period?

<!-- grade_id: warm-times, points: 0, grade: false, solution: true -->
```python
# Edit the variable to find correct answer
warm_times = 'XXX'

### BEGIN SOLUTION
warm_times = len(warm)
### END SOLUTION
```

<!-- grade_id: warm-times-test, points: 0.33, grade: true, solution: false -->
```python
# Print the answer
print("In April 2013 it was warm", warm_times, "times.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
# It was slippery 16 times
assert_equal(warm_times, 0)

### END HIDDEN TESTS
```

- How many times was it cold during the study period?

<!-- grade_id: cold-times, points: 0, grade: false, solution: true -->
```python
# Edit the variable to find correct answer
cold_times = 'XXX'

### BEGIN SOLUTION
cold_times = len(cold)
### END SOLUTION
```

<!-- grade_id: cold-times-test, points: 0.34, grade: true, solution: false -->
```python
# Print the answer
print("In April 2013 it was cold", cold_times, "times.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
# It was slippery 16 times
assert_equal(cold_times, 15)

### END HIDDEN TESTS
```

### Comment section:

Here, we ask a few questions to make sure you have understand the concepts etc. Answer shortly with a couple of sentences to following questions.

You can also write any feedback or questions concerning this problem (we will also use this space for giving feedback to you, if necessary).

#### Questions

1: Is the concept of conditional statements clear to you? If not, what is difficult to understand?

2: Did you include comments in your code blocks?
  
Write your answers below:

<!-- grade_id: p2-style-assessment, points: 1, grade: true, solution: true -->


## Problem 3 - Allocating locations (**8 points**)


### Overview 
Following map shows the locations of the weather stations (as blue points) in Finland that are more than 70 years old [1].
In this problem we are interested to find out whether the station network was equally distributed across Finland
seventy years ago. We have divided Finland into four geographical zones (i.e. North West, North East, South West, South East)
according the approximate center point of Finnish mainland located at `26.3, 64.5` (lon-lat in decimal degrees).

![](img/FMI_stations_70_years_older.png)

[1]: The locations and the age of weather stations were obtained from: http://en.ilmatieteenlaitos.fi/observation-stations

Below, we have given you the coordinates of 34 weather stations.
The location of a single station is determined with a pair of latitude and longitude coordinates.
The coordinates of all the stations are separated into two lists (`lat` and `lon`) and the names of the stations are in `stations` list. From these lists, you would get e.g. the location of the first station by combining the latitude and longitude coordinates from coordinate lists, and the name of that station from `stations` list at index[0]:

### Problem statement

In this problem your job is to print the names of weather stations located in different zones. **Optionally** you should also report the share of weather stations that allocated to each zone that could be used to evaluate if certain zone was over/under-represented seventy years ago (optional task, does not affect on grading).

To solve this problem, you should modify and fill in the missing parts in the code cells below.

**Your score on this problem will be based on following criteria:**

 1. Create four lists for geographical zones in Finland (i.e. `nort_west`, `north_east`, `south_west`, `south_east`)
 2. Iterate over values and determine to which geographical zone the station belongs
    1. Hint: You should create a loop that iterates `N` -number of times. Create a variable *`N`* that should contain the number of stations we have here.
    2. You should use a conditional statement to find out if the latitude coordinate of a station is either North or South of the center point of Finland (`26.3, 64.5`) **AND** if the longitude location is West or East from that center point.
    3. You should insert the name of the station into the correct geographical zone list (step 1)
 3. Print out the names of stations at each geographical zone
 4. **Optional:** Calculate and print the share of stations at each zone (the total number of stations equals to 100 %)

### Data

Here, we provide you the data you should use to solve the problem 3 (run the cell to get values into the memory of your computer):

```python
# Station names
stations = ['Hanko Russarö', 'Heinola Asemantaus', 'Helsinki Kaisaniemi', 
            'Helsinki Malmi airfield', 'Hyvinkää Hyvinkäänkylä', 'Joutsa Savenaho', 
            'Juuka Niemelä', 'Jyväskylä airport', 'Kaarina Yltöinen', 'Kauhava airfield', 
            'Kemi Kemi-Tornio airport', 'Kotka Rankki', 'Kouvola Anjala', 
            'Kouvola Utti airport', 'Kuopio Maaninka', 'Kuusamo airport', 
            'Lieksa Lampela', 'Mustasaari Valassaaret', 'Parainen Utö', 'Pori airport', 
            'Rovaniemi Apukka', 'Salo Kärkkä', 'Savonlinna Punkaharju Laukansaari', 
            'Seinäjoki Pelmaa', 'Siikajoki Ruukki', 'Siilinjärvi Kuopio airport', 
            'Tohmajärvi Kemie', 'Utsjoki Nuorgam', 'Vaala Pelso', 'Vaasa airport', 
            'Vesanto Sonkari', 'Vieremä Kaarakkala', 'Vihti Maasoja', 'Ylitornio Meltosjärvi']

# Latitude coordinates of Weather stations  
lats = [59.77, 61.2, 60.18, 60.25, 60.6, 61.88, 63.23, 62.4,
       60.39, 63.12, 65.78, 60.38, 60.7, 60.9, 63.14, 65.99,
       63.32, 63.44, 59.78, 61.47, 66.58, 60.37, 61.8, 62.94,
       64.68, 63.01, 62.24, 70.08, 64.5, 63.06, 62.92, 63.84,
       60.42, 66.53]

# Longitude coordinates of Weather stations 
lons = [22.95, 26.05, 24.94, 25.05, 24.8, 26.09, 29.23, 25.67, 
       22.55, 23.04, 24.58, 26.96, 26.81, 26.95, 27.31, 29.23, 
       30.05, 21.07, 21.37, 21.79, 26.01, 23.11, 29.32, 22.49, 
       25.09, 27.8, 30.35, 27.9, 26.42, 21.75, 26.42, 27.22, 
       24.4, 24.65]

# Cutoff values that correspond to the centroid of Finnish mainland
# North - South
north_south_cutoff = 64.5

# East-West
east_west_cutoff = 26.3
```

- Create four empty lists for geographical zones in Finland, use **exactly** following variable names (**0.5 points**)
   - `north_west`
   - `north_east`
   - `south_west`
   - `south_east`

<!-- grade_id: geo-zones, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# Empty lists for stations located at different geographical zones
north_west = []
north_east = []
south_west = []
south_east = []
### END SOLUTION
```

<!-- grade_id: geo-zones-test, points: 0.5, grade: true, solution: false -->
```python
# This test print should work
print(north_west, south_west)

### BEGIN HIDDEN TESTS
from nose.tools import ok_
# Test that variables exist
ok_("north_west" in locals())
ok_("north_east" in locals())
ok_("south_west" in locals())
ok_("south_east" in locals())

# Test that they are lists
ok_(isinstance(north_west, list) )
ok_(isinstance(north_east, list) )
ok_(isinstance(south_west, list) )
ok_(isinstance(south_east, list) )

# Test that they are empty
ok_(len(north_west) == 0 )
ok_(len(north_east) == 0 )
ok_(len(south_west) == 0 )
ok_(len(south_east) == 0 )

### END HIDDEN TESTS
```

**Step 1**: Count the number of stations and store that value in variable: *`N`* (*0.5 points*)

<!-- grade_id: station_count, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# Determine the number of stations that we have in total
N = len(stations)
### END SOLUTION
```

<!-- grade_id: station_count_test, points: 0.5, grade: true, solution: false -->
```python
# How many stations do we have?
print("In the data, there are", N, "stations.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(N, 34)
### END HIDDEN TESTS
```

**Step 2**: Make a loop that iterates N number of times, and allocates stations to different geographical zones based on the coordinates. 

<!-- grade_id: allocate_stations, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# Step 2. Make a loop that iterates N number of times. 
# N should be based on the number of stations that we have here. 
# --------------------------------------------------------------

# Create a for loop that runs N number of times
for i in range(N):
    
    # Step 2.1 Find out the latitude, longitude and the name of the station at index i
    lat = lats[i]
    lon = lons[i]
    station = stations[i]
    
    # Steps 2.2-2.3 - Make conditional statements to add the station to correct zone -list
    # You should determine if 
    # a) the lat coordinate is North or South of the center latitude (stored in NScutoff variable)
    # AND 
    # b) if the lon coordinate is East or West of the center longitude (stored in EWcutoff variable)
    # c) Based on the criteria you should add the NAME of the station (from stations -list) to correct list (done in step 1)
    
    # Make conditional statements to allocate the stations to correct zones
    
    # Check if the station is South from central point
    if lat <= north_south_cutoff:
        # Check if the station is in West
        if lon <= east_west_cutoff:
            # If it was South-West add the station name to corresponding list    
            south_west.append(station)
        else:
            # If it was not, it must be located in the South-East part of Finland
            south_east.append(station)
    # Otherwise the station must be North (hence we can use else)
    else:
        # Check if the station is in West
        if lon <= east_west_cutoff:
            # If it was West, add the station name to North-West list    
            north_west.append(station)
        else:
            # If it was not, it must be located in the North-East part of Finland
            north_east.append(station)
            
### END SOLUTION
```

**Step 3:** In the following cells, you should print out the correct names for each geographical zone. (*1.5 points each - total: 6 points*)
    
   - We have given you the correct number of stations as a hint for each zone, so it is easier for you to know whether you have correct answer. 

<!-- grade_id: allocate_stations_test_nw, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in North West
# Hint: there should be 4 stations in this class
print("The names of the North-West stations are:\n", north_west)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
nw_correct = ['Kemi Kemi-Tornio airport', 'Rovaniemi Apukka', 'Siikajoki Ruukki', 'Ylitornio Meltosjärvi']

# Test the results
assert_equal(north_west, nw_correct)

### END HIDDEN TESTS
```

<!-- grade_id: allocate_stations_test_ne, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in North Eest
# Hint: there should be 2 stations in this class
print("The names of the North-East stations are:\n", north_east)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
ne_correct = ['Kuusamo airport', 'Utsjoki Nuorgam']

# Test the results
assert_equal(north_east, ne_correct)

### END HIDDEN TESTS
```

<!-- grade_id: allocate_stations_test_sw, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in South West
# Hint: there should be 16 stations in this class
print("The names of the South-West stations are:\n", south_west)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
sw_correct = ['Hanko Russarö', 'Heinola Asemantaus', 'Helsinki Kaisaniemi', 'Helsinki Malmi airfield', 'Hyvinkää Hyvinkäänkylä', 'Joutsa Savenaho', 'Jyväskylä airport', 'Kaarina Yltöinen', 'Kauhava airfield', 'Mustasaari Valassaaret', 'Parainen Utö', 'Pori airport', 'Salo Kärkkä', 'Seinäjoki Pelmaa', 'Vaasa airport', 'Vihti Maasoja']

# Test the results
assert_equal(south_west, sw_correct)

### END HIDDEN TESTS
```

<!-- grade_id: allocate_stations_test_se, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in South East
# Hint: there should be 12 stations in this class
print("The names of the South-East stations are:\n", south_east)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
se_correct = ['Juuka Niemelä', 'Kotka Rankki', 'Kouvola Anjala', 'Kouvola Utti airport', 'Kuopio Maaninka', 'Lieksa Lampela', 'Savonlinna Punkaharju Laukansaari', 'Siilinjärvi Kuopio airport', 'Tohmajärvi Kemie', 'Vaala Pelso', 'Vesanto Sonkari', 'Vieremä Kaarakkala']
assert_equal(south_east, se_correct)

### END HIDDEN TESTS
```

### Comment section:

Here, we ask a few questions to make sure you have understand the concepts etc. Answer shortly with a couple of sentences to following questions.


You can also write any feedback or questions concerning this problem (we will also use this space for giving feedback to you, if necessary).

#### Questions

1: Is the concept of using **AND** operator and the difference between **if - elif - else** conditional statements clear to you? If not, what is difficult to understand?

2: Did you include comments in your code blocks?

Write your answers below:

<!-- grade_id: p3-style-assessment, points: 1, grade: true, solution: true -->


**Step 4 (optional)**: Print the share of stations at each geographical zone (in percentages) (*does not affect grading*)

 - Store the answers into variables:
    - `north_west_share`
    - `north_east_share`
    - `south_west_share`
    - `south_east_share`


<!-- grade_id: station-percentage, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# We can calculate the shares of each zone by dividing the number of stations
# in each zone divided by the total number of stations (stored in variable N).
# Multiplying by hundred gives the values in percentages (valuerange 0-100)
north_west_share = (len(north_west) / N) * 100
north_east_share = (len(north_east) / N) * 100
south_west_share = (len(south_west) / N) * 100
south_east_share = (len(south_east) / N) * 100

### END SOLUTION
```

- Print the results following format: "North-West contains 99 % of all stations." (this is an example not a correct answer)

<!-- grade_id: station-percentage-tests, points: 0, grade: true, solution: false -->
```python
# Print the information (you don't need to modify this)
# .format() is a Python function that can be used to easily insert values inside a text-template such as below.
# .0f below is a specific operator that rounds the decimal values into whole numbers
print("North-West contains{share: .0f} % of all stations.".format(share=north_west_share))
print("North-East contains{share: .0f} % of all stations.".format(share=north_east_share))
print("South-West contains{share: .0f} % of all stations.".format(share=south_west_share))
print("South-East contains{share: .0f} % of all stations.".format(share=south_east_share))

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal

# Test if variables exist
ok_('north_west_share' in locals())
ok_('north_east_share' in locals())
ok_('south_west_share' in locals())
ok_('south_east_share' in locals())

# Test correct values (here truncating values is enough to ensure correct values)
assert_equal(int(north_west_share), 11)
assert_equal(int(north_east_share), 5)
assert_equal(int(south_west_share), 47)
assert_equal(int(south_east_share), 35)
### END HIDDEN TESTS
```

## Extra Problem - Nested `for` loops (optional, does not affect grading)

**Warning:** This is a difficult problem. Try to do this task only if you have confidence in programming, and you are up for a little challenge. :) 

In addition to having single `for` loops that iterate across some variable range, it is possible to *nest* `for` loops within one another.

Consider the example below:

```python
>>> for char in 'dog':
...     for char2 in 'cat':
...         print (char, char2)
    d c
    d a
    d t
    o c
    o a
    o t
    g c
    g a
    g t
```

Here, you can see that in the first pass through the first `for` loop, the value of `char` is `d`.
Entering the inner (or nested) loop, `char2` is set to `c`.
After this, the output is written to the screen and since there are more letters to loop over in the inner `for` loop, the value of `char2` will be updated upon the next pass.
The second time through the inner loop, `char2` is set to `a` while `char` remains `d`.
Like this, the inner loop will run through all of the letters in `cat` for each time that `char` is updated in the outer loop.
It doesn't take too much imagination to figure out this is a very useful concept.

For this problem you should do following:

1. Create a variable `star` with text `"*"` and an empty string variable `text`. Recall, you can created empty string variables by assigning `""` as their value.
2. Use nested `for` loops and the variables above to produce the text formation below when `print(text)` is run at the end of your script.

    ```
    *******
    *******
    *******
    ```
3. Create a variable `line` with text `"-"` and an empty string variable `flag`.
4. Use nested `for` loops and the variables above to produce the text formation below when `print(flag)` is run at the end of your script. **Note**: You will need to use conditional statements to produce the desired output.

    ```
    *******------------
    *******------------
    *******------------
    -------------------
    -------------------
    ```


- Your code for producing the asterix square:

<!-- grade_id: asterix-square, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
star = '*'
text = ""

for i in range(3):
    for j in range(7):
        text+= star
    text += '\n'

### END SOLUTION
text
```

<!-- grade_id: asterix-square-test, points: 0, grade: true, solution: false -->
```python
print(text)
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
correct_shape = '*******\n*******\n*******\n'
assert_equal(text, correct_shape)
### END HIDDEN TESTS
```

- Your code for producing the flag:

<!-- grade_id: flag, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
star = '*'
flag = ""
line = '-'

for i in range(5):
    # For the first three lines add asterix and lines
    if i < 3:
        for j in range(7):
            flag+= star
            
        for k in range(12):
            flag+=line
    # For the last two lines add only lines
    else:
        for k in range(19):
            flag+=line
    flag += '\n'
### END SOLUTION
```

<!-- grade_id: flag-test, points: 0, grade: true, solution: false -->
```python
print(flag)
### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
correct_shape = '*******------------\n*******------------\n*******------------\n-------------------\n-------------------\n'
assert_equal(flag, correct_shape)
### END HIDDEN TESTS
```