## Problem 3 - Allocating locations (*8 points*)


### Overview 
Following map shows the locations of the weather stations (as blue points) in Finland that are more than 70 years old [1].
In this problem we are interested to find out whether the station network was equally distributed across Finland
seventy years ago. We have divided Finland into four geographical zones (i.e. North West, North East, South West, South East)
according the approximate center point of Finnish mainland located at `26.3, 64.5` (lon-lat in decimal degrees).

![](img/FMI_stations_70_years_older.png)

[1]: The locations and the age of weather stations were obtained from: http://en.ilmatieteenlaitos.fi/observation-stations

Below, we have given you the coordinates of 34 weather stations.
The location of a single station is determined with a pair of latitude and longitude coordinates.
The coordinates of all the stations are separated into two lists (`lat` and `lon`) and the names of the stations are in `stations` list. From these lists, you would get e.g. the location of the first station by combining the latitude and longitude coordinates from coordinate lists, and the name of that station from `stations` list at index[0]:

### Problem statement

In this problem your job is to print the names of weather stations located in different zones. **Optionally** you should also report the share of weather stations that allocated to each zone that could be used to evaluate if certain zone was over/under-represented seventy years ago (optional task, does not affect on grading).

To solve this problem, you should modify and fill in the missing parts in the code cells below.

**Notice**: follow closely the instructions, e.g. using **exactly** the same variable names as mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting as in the instructions. 


**Your score on this problem will be based on following criteria:**

 1. Create four lists for geographical zones in Finland (i.e. `nort_west`, `north_east`, `south_west`, `south_east`)
 2. Iterate over values and determine to which geographical zone the station belongs
    1. Hint: You should create a loop that iterates `N` -number of times. Create a variable *`N`* that should contain the number of stations we have here.
    2. You should use a conditional statement to find out if the latitude coordinate of a station is either North or South of the center point of Finland (`26.3, 64.5`) **AND** if the longitude location is West or East from that center point.
    3. You should insert the name of the station into the correct geographical zone list (step 1)
 3. Print out the names of stations at each geographical zone
 4. **Optional:** Calculate and print the share of stations at each zone (the total number of stations equals to 100 %)

### Data

Here, we provide you the data you should use to solve the problem 3 (run the cell to get values into the memory of your computer):

```python
# Station names
stations = ['Hanko Russarö', 'Heinola Asemantaus', 'Helsinki Kaisaniemi', 
            'Helsinki Malmi airfield', 'Hyvinkää Hyvinkäänkylä', 'Joutsa Savenaho', 
            'Juuka Niemelä', 'Jyväskylä airport', 'Kaarina Yltöinen', 'Kauhava airfield', 
            'Kemi Kemi-Tornio airport', 'Kotka Rankki', 'Kouvola Anjala', 
            'Kouvola Utti airport', 'Kuopio Maaninka', 'Kuusamo airport', 
            'Lieksa Lampela', 'Mustasaari Valassaaret', 'Parainen Utö', 'Pori airport', 
            'Rovaniemi Apukka', 'Salo Kärkkä', 'Savonlinna Punkaharju Laukansaari', 
            'Seinäjoki Pelmaa', 'Siikajoki Ruukki', 'Siilinjärvi Kuopio airport', 
            'Tohmajärvi Kemie', 'Utsjoki Nuorgam', 'Vaala Pelso', 'Vaasa airport', 
            'Vesanto Sonkari', 'Vieremä Kaarakkala', 'Vihti Maasoja', 'Ylitornio Meltosjärvi']

# Latitude coordinates of Weather stations  
lats = [59.77, 61.2, 60.18, 60.25, 60.6, 61.88, 63.23, 62.4,
       60.39, 63.12, 65.78, 60.38, 60.7, 60.9, 63.14, 65.99,
       63.32, 63.44, 59.78, 61.47, 66.58, 60.37, 61.8, 62.94,
       64.68, 63.01, 62.24, 70.08, 64.5, 63.06, 62.92, 63.84,
       60.42, 66.53]

# Longitude coordinates of Weather stations 
lons = [22.95, 26.05, 24.94, 25.05, 24.8, 26.09, 29.23, 25.67, 
       22.55, 23.04, 24.58, 26.96, 26.81, 26.95, 27.31, 29.23, 
       30.05, 21.07, 21.37, 21.79, 26.01, 23.11, 29.32, 22.49, 
       25.09, 27.8, 30.35, 27.9, 26.42, 21.75, 26.42, 27.22, 
       24.4, 24.65]

# Cutoff values that correspond to the centroid of Finnish mainland
# North - South
north_south_cutoff = 64.5

# East-West
east_west_cutoff = 26.3
```

- Create four empty lists for geographical zones in Finland, use **exactly** following variable names (**0.5 points**)
   - `north_west`
   - `north_east`
   - `south_west`
   - `south_east`

<!-- grade_id: geo-zones, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# Empty lists for stations located at different geographical zones
north_west = []
north_east = []
south_west = []
south_east = []
### END SOLUTION
```

<!-- grade_id: geo-zones-test, points: 0.5, grade: true, solution: false -->
```python
# This test print should work
print(north_west, south_west)

### BEGIN HIDDEN TESTS
from nose.tools import ok_
# Test that variables exist
ok_("north_west" in locals())
ok_("north_east" in locals())
ok_("south_west" in locals())
ok_("south_east" in locals())

# Test that they are lists
ok_(isinstance(north_west, list) )
ok_(isinstance(north_east, list) )
ok_(isinstance(south_west, list) )
ok_(isinstance(south_east, list) )

# Test that they are empty
ok_(len(north_west) == 0 )
ok_(len(north_east) == 0 )
ok_(len(south_west) == 0 )
ok_(len(south_east) == 0 )

### END HIDDEN TESTS
```

**Step 1**: Count the number of stations and store that value in variable: *`N`* (*0.5 points*)

<!-- grade_id: station_count, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# Determine the number of stations that we have in total
N = len(stations)
### END SOLUTION
```

<!-- grade_id: station_count_test, points: 0.5, grade: true, solution: false -->
```python
# How many stations do we have?
print("In the data, there are", N, "stations.")

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal
assert_equal(N, 34)
### END HIDDEN TESTS
```

**Step 2**: Make a loop that iterates N number of times, and allocates stations to different geographical zones based on the coordinates. 

<!-- grade_id: allocate_stations, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# Step 2. Make a loop that iterates N number of times. 
# N should be based on the number of stations that we have here. 
# --------------------------------------------------------------

# Create a for loop that runs N number of times
for i in range(N):
    
    # Step 2.1 Find out the latitude, longitude and the name of the station at index i
    lat = lats[i]
    lon = lons[i]
    station = stations[i]
    
    # Steps 2.2-2.3 - Make conditional statements to add the station to correct zone -list
    # You should determine if 
    # a) the lat coordinate is North or South of the center latitude (stored in NScutoff variable)
    # AND 
    # b) if the lon coordinate is East or West of the center longitude (stored in EWcutoff variable)
    # c) Based on the criteria you should add the NAME of the station (from stations -list) to correct list (done in step 1)
    
    # Make conditional statements to allocate the stations to correct zones
    
    # Check if the station is South from central point
    if lat <= north_south_cutoff:
        # Check if the station is in West
        if lon <= east_west_cutoff:
            # If it was South-West add the station name to corresponding list    
            south_west.append(station)
        else:
            # If it was not, it must be located in the South-East part of Finland
            south_east.append(station)
    # Otherwise the station must be North (hence we can use else)
    else:
        # Check if the station is in West
        if lon <= east_west_cutoff:
            # If it was West, add the station name to North-West list    
            north_west.append(station)
        else:
            # If it was not, it must be located in the North-East part of Finland
            north_east.append(station)
            
### END SOLUTION
```

**Step 3:** In the following cells, you should print out the correct names for each geographical zone. (*1.5 points each - total: 6 points*)
    
   - We have given you the correct number of stations as a hint for each zone, so it is easier for you to know whether you have correct answer. 

<!-- grade_id: allocate_stations_test_nw, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in North West
# Hint: there should be 4 stations in this class
print("The names of the North-West stations are:\n", north_west)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
nw_correct = ['Kemi Kemi-Tornio airport', 'Rovaniemi Apukka', 'Siikajoki Ruukki', 'Ylitornio Meltosjärvi']

# Test the results
assert_equal(north_west, nw_correct)

### END HIDDEN TESTS
```

<!-- grade_id: allocate_stations_test_ne, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in North Eest
# Hint: there should be 2 stations in this class
print("The names of the North-East stations are:\n", north_east)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
ne_correct = ['Kuusamo airport', 'Utsjoki Nuorgam']

# Test the results
assert_equal(north_east, ne_correct)

### END HIDDEN TESTS
```

<!-- grade_id: allocate_stations_test_sw, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in South West
# Hint: there should be 16 stations in this class
print("The names of the South-West stations are:\n", south_west)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
sw_correct = ['Hanko Russarö', 'Heinola Asemantaus', 'Helsinki Kaisaniemi', 'Helsinki Malmi airfield', 'Hyvinkää Hyvinkäänkylä', 'Joutsa Savenaho', 'Jyväskylä airport', 'Kaarina Yltöinen', 'Kauhava airfield', 'Mustasaari Valassaaret', 'Parainen Utö', 'Pori airport', 'Salo Kärkkä', 'Seinäjoki Pelmaa', 'Vaasa airport', 'Vihti Maasoja']

# Test the results
assert_equal(south_west, sw_correct)

### END HIDDEN TESTS
```

<!-- grade_id: allocate_stations_test_se, points: 1.5, grade: true, solution: false -->
```python
# This test print should print out station names in South East
# Hint: there should be 12 stations in this class
print("The names of the South-East stations are:\n", south_east)

### BEGIN HIDDEN TESTS
from nose.tools import assert_equal

# Correctly allocated stations
se_correct = ['Juuka Niemelä', 'Kotka Rankki', 'Kouvola Anjala', 'Kouvola Utti airport', 'Kuopio Maaninka', 'Lieksa Lampela', 'Savonlinna Punkaharju Laukansaari', 'Siilinjärvi Kuopio airport', 'Tohmajärvi Kemie', 'Vaala Pelso', 'Vesanto Sonkari', 'Vieremä Kaarakkala']
assert_equal(south_east, se_correct)

### END HIDDEN TESTS
```

### Comment section:

Here, we ask a few questions to make sure you have understand the concepts etc. Answer shortly with a couple of sentences to following questions.


You can also write any feedback or questions concerning this problem (we will also use this space for giving feedback to you, if necessary).

#### Questions

1: Is the concept of using **AND** operator and the difference between **if - elif - else** conditional statements clear to you? If not, what is difficult to understand?

2: Did you include comments in your code blocks?

Write your answers below (double click to activate the cell below):

<!-- grade_id: p3-style-assessment, points: 1, grade: true, solution: true -->


**Step 4 (optional)**: Print the share of stations at each geographical zone (in percentages) (*does not affect grading*)

 - Store the answers into variables:
    - `north_west_share`
    - `north_east_share`
    - `south_west_share`
    - `south_east_share`


<!-- grade_id: station-percentage, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# We can calculate the shares of each zone by dividing the number of stations
# in each zone divided by the total number of stations (stored in variable N).
# Multiplying by hundred gives the values in percentages (valuerange 0-100)
north_west_share = (len(north_west) / N) * 100
north_east_share = (len(north_east) / N) * 100
south_west_share = (len(south_west) / N) * 100
south_east_share = (len(south_east) / N) * 100

### END SOLUTION
```

- Print the results following format: "North-West contains 99 % of all stations." (this is an example not a correct answer)

<!-- grade_id: station-percentage-tests, points: 0, grade: true, solution: false -->
```python
# Print the information (you don't need to modify this)
# .format() is a Python function that can be used to easily insert values inside a text-template such as below.
# .0f below is a specific operator that rounds the decimal values into whole numbers
print("North-West contains{share: .0f} % of all stations.".format(share=north_west_share))
print("North-East contains{share: .0f} % of all stations.".format(share=north_east_share))
print("South-West contains{share: .0f} % of all stations.".format(share=south_west_share))
print("South-East contains{share: .0f} % of all stations.".format(share=south_east_share))

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal

# Test if variables exist
ok_('north_west_share' in locals())
ok_('north_east_share' in locals())
ok_('south_west_share' in locals())
ok_('south_east_share' in locals())

# Test correct values (here truncating values is enough to ensure correct values)
assert_equal(int(north_west_share), 11)
assert_equal(int(north_east_share), 5)
assert_equal(int(south_west_share), 47)
assert_equal(int(south_east_share), 35)
### END HIDDEN TESTS
```

#### Done!

That's it! Now you have finished all the assignments for this week!

If you want, you can still continue with optional [Problem 4](Exercise-3-problem-4.ipynb).