## Before starting

Test that your Python version is up-to-date. Our automatic grading system requires to have IPython version which is higher than 3. Test your version by running following cell (shift + Enter):

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it. Ask help, if you don't know how."
print("Your system is okay, you can continue! :) ")
```

## Problem 1 - Batch processing data files with a `for` loop (*4 points*)

This problem is meant to simulate a common problem dealing with data files: batch processing.

Batch processing involves using scripts to repeat a process with many data files, and one common task is generating a list of filenames that will be processed or saved to disk.

For this problem you will need to modify the code blocks below and add code to produce the desired outcomes. 

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria:**

 - Creating and using variables to produce the desired text format
 - Successfully using a for loop to iterate over desired set of numbers
 - Successfully producing the desired filename(s)
 - Including comments that explain what most lines in the code do
 - Answering a couple questions at the end of the problem
 - Uploading your notebook to your GitHub repository for this week's exercise

### Part 1 (*0.25 points*)

- Create a new variable called `basename` that contains text `"Station"`.

<!-- grade_id: station-name, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem-1-station-name-test, points: 0.25, grade: true, solution: false -->
```python
# Test print of the variable
print(basename)

```

### Part 2 (*0.25 points*)

- Create a new variable `filenames` that is an empty list.

<!-- grade_id: empty-filenames-list, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem-1-empty-filenames-list-test, points: 0.25, grade: true, solution: false -->
```python
# Test print of the variable
print(filenames)

```

### Part 3 (*2.5 points*)

- Iterate over the number range 0-20 and
    - Create a variable `station` that contains the 1) text from `basename` variable, 2) the number, and 3) the file extension `.txt`  
    - Add the content of `station` to `filenames` list which should have following content in the end:

      ```
      ['Station_0.txt', 'Station_1.txt', 'Station_2.txt', 'Station_3.txt',
       'Station_4.txt', 'Station_5.txt', 'Station_6.txt', 'Station_7.txt',
       'Station_8.txt', 'Station_9.txt', 'Station_10.txt', 'Station_11.txt',
       'Station_12.txt', 'Station_13.txt', 'Station_14.txt', 'Station_15.txt',
       'Station_16.txt', 'Station_17.txt', 'Station_18.txt', 'Station_19.txt',
       'Station_20.txt']
      ```

<!-- grade_id: populate-filenames-list, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem-1-populate-filenames-list-test1, points: 0.5, grade: true, solution: false -->
```python
# Test print of the results
print(filenames)

```

<!-- grade_id: problem-1-populate-filenames-list-test2, points: 0.5, grade: true, solution: false -->
```python
from nose.tools import ok_, assert_equal
# Check that the value of the last station is correct
assert_equal(station.lower().strip(), "station_20.txt", 'The value of the last station is not correct')
```

<!-- grade_id: problem-1-populate-filenames-list-test3, points: 0.5, grade: true, solution: false -->
```python
from nose.tools import ok_, assert_equal
# Check that there exists 21 values in the list
ok_(len(filenames) == 21, 'The length of the list "filenames" is not 21')
```

<!-- grade_id: problem-1-populate-filenames-list-test4, points: 1, grade: true, solution: false -->
```python
# Check that the content of the list is correct
```

### Part 4 (*1 point*)

Here, we ask a few questions to make sure you have understand the concepts, etc. Answer shortly with a few sentences in the Markdown cell below.

You can also write any feedback or questions concerning this problem in this Markdown cell.

1. Is the concept of a loop clear to you? If not, what is difficult to understand?
2. Did you include comments in your code blocks?
  
Write your answers below (double click to activate and edit the cell below).

<!-- grade_id: problem-1-p1-style-assessment, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

### Done!

That's it! Now you are ready to continue with [Problem 2](Exercise-3-problem-2.ipynb).