# Exercise 3

The exercise for this week is meant to help you to understand `for` loops and conditional statements in Python.
Below you have a series of problems in which you are asked to edit the notebook files and write the code necessary to produce the desired results.

After making your changes to the Notebook, you will need to upload it to GitHub.


- **Exercise 3 is due by the start of lecture on 26.9**.

- Don't forget to check out the [hints for this week's exercise](https://geo-python.github.io/2018/lessons/L3/exercise-3-hints.html) if you're having trouble.

- Scores on this exercise are out of **20 points**.

## Before starting

Test that your Python version is up-to-date. Our automatic grading system requires to have IPython version which is higher than 3. Test your version by running following cell (shift + Enter):

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
print("Your system is okay, you can continue! :) ")
```

## Problem 1 - Batch processing data files with a `for` loop (**4 Points**)

This problem is meant to simulate a common problem dealing with data files: batch processing.

Batch processing involves using scripts to process many data files, and one common task is generating a list of filenames that will be processed, or saved to disk.

For this problem you will need to modify the code blocks below and add codes to produce the desired outcomes. 

**Notice**: follow closely the instructions, e.g. using **exactly** the same variable names as mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting as in the instructions. 


1: Create a new variable called `basename` that contains text `"Station"`.

<!-- grade_id: station-name, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
basename = "Station"
### END SOLUTION
```

<!-- grade_id: station-name-correct, points: 0.25, grade: true, solution: false -->
```python
# Test print of the variable
print(basename)

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal
# Validate variable name
ok_('basename' in locals())

# Validate value
assert_equal(basename, 'Station')
### END HIDDEN TESTS
```

2: Create a new variable `filenames` that is an empty list.

<!-- grade_id: empty-filenames-list, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
filenames = []
### END SOLUTION
```

<!-- grade_id: empty-filenames-list-correct, points: 0.25, grade: true, solution: false -->
```python
# Test print of the variable
print(filenames)

### BEGIN HIDDEN TESTS

from nose.tools import ok_
ok_('filenames' in locals())
ok_(isinstance(filenames, list))

### END HIDDEN TESTS
```

3: Iterate over the number range 0-20 and
  - Create a variable `station` that contains the 1) text from `basename` variable, 2) the number, and 3) the file extension `.txt`  
  - Add the content of `station` to `filenames` list which should have following content in the end:

      ```
      ['Station_0.txt', 'Station_1.txt', 'Station_2.txt', 'Station_3.txt',
       'Station_4.txt', 'Station_5.txt', 'Station_6.txt', 'Station_7.txt',
       'Station_8.txt', 'Station_9.txt', 'Station_10.txt', 'Station_11.txt',
       'Station_12.txt', 'Station_13.txt', 'Station_14.txt', 'Station_15.txt',
       'Station_16.txt', 'Station_17.txt', 'Station_18.txt', 'Station_19.txt',
       'Station_20.txt']
      ```

<!-- grade_id: populate-filenames-list, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# Iterate from 0 up to 20
for i in range(21):
    # Create the name
    station = basename + "_" + str(i) + ".txt"
    # Append to lis
    filenames.append(station)

# Print the contents
print(filenames)
    
### END SOLUTION
```

<!-- grade_id: populate-filenames-list-1, points: 3.5, grade: true, solution: false -->
```python
# Test print of the results
print(filenames)

### BEGIN HIDDEN TESTS
from nose.tools import ok_, assert_equal

# Check that 'station' variable exists
ok_('station' in locals())

# Check that the value of the last station is correct
assert_equal(station.lower().strip(), "station_20.txt")

# Check that there exists 21 values in the list
ok_(len(filenames) == 21)

correct_list = ['station_0.txt', 'station_1.txt', 'station_2.txt', 'station_3.txt', 'station_4.txt', 'station_5.txt', 'station_6.txt', 'station_7.txt', 'station_8.txt', 'station_9.txt', 'station_10.txt', 'station_11.txt', 'station_12.txt', 'station_13.txt', 'station_14.txt', 'station_15.txt', 'station_16.txt', 'station_17.txt', 'station_18.txt', 'station_19.txt', 'station_20.txt']

# Check that the contents match (use lowercase to allow some variation in upper/lower case)
for value in filenames:
    ok_(value.strip().lower() in correct_list)

### END HIDDEN TESTS
```